evaluation_interval: 1m

rule_files:
  - prometheus.always_firing.yaml

tests:
  - interval: 1m
    input_series: []
    alert_rule_test:
      - eval_time: 2m
        alertname: AlwaysFiring
        exp_alerts:
          - exp_labels:
              severity: info
              source_cluster: no-cluster
            exp_annotations:
              summary: Test alert - please ignore
              description: Test alert description - please ignore
